## Assignment 3 (Due 11 November 2016 by 11:59pm)

---
## Part a (5 points)
Implement an object-oriented PRNG based on the SHA-256 hash of a seed concatenated with a sequence number, in Python, by subclassing the class Random and overwriting the methods random(), seed(), getstate(), setstate(), jumpahead(), and getrandbits(). Your implementation should inherit shuffle(), choice(), sample(), randbelow(), etc., from the parent Random class.

<!-- grade_id: a_random, points: 4, grade: true, solution: true -->
```python
import random
import hashlib as _hashlib
from os import urandom as _urandom
from binascii import hexlify as _hexlify

class FancyRandom(random.Random):
    """
    PRNG based on the SHA-256 hash of a seed concatenated with a sequence number
    You should overwrite the methods random(), seed(), getstate(), setstate(), jumpahead(), and getrandbits()
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: a_test1, points: 1, grade: true, solution: false -->
```python
def test_random01(seed):
    """
    Test the random number is in the range [0.0, 1.0)
    """
    random_generator = FancyRandom()
    random_generator.seed(seed)
    res = random_generator.random()
    return res < 1 and res >= 0 

assert test_random01(123456)

def test_random01_all(n):
    """
    Test the random number is in the range [0.0, 1.0)
    for n different seeds
    """
    for i in range(n):
        if not test_random01(i):
            return False
    return True

assert test_random01_all(1000)
```

---
## Part b (8 points)
Implement tests of: 

(at least one for each, more tests will be **bonus**)
1. setting the seed: at least one test to take care of int, long, str seed
2. setting the state
3. getting the state
4. jumping ahead: at least one test for transitivity, "jump(a+b) = jump(a) followed by jump(b)"
5. getrandbits: at least one test to check exactly k bits are given
6. at least two tests for uniformity, with “reasonable” power to detect a bug, including a test using the Kolmogorov-Smirnov statistic for random() and a binomial test for the single-bit frequency for getrandbits().
7. tests that your class inherited the expected methods from the parent class

B.1. setting the seed: at least one test to check same seed gives same result

<!-- grade_id: b_1_test, points: 1, grade: true, solution: true -->
```python
def test_seed():
    """
    at least check same seed will give same result
    """
    # YOUR CODE HERE
    raise NotImplementedError()
    assert a == b
    
test_seed()
```

B.2. setting the state

<!-- grade_id: b_2_test, points: 1, grade: true, solution: true -->
```python
def test_setstate():
    # YOUR CODE HERE
    raise NotImplementedError()
    
test_setstate()
```

B.3. getting the state

<!-- grade_id: b_3_test, points: 1, grade: true, solution: true -->
```python
def test_getstate():
    # YOUR CODE HERE
    raise NotImplementedError()
    
test_getstate()
```

B.4. jumping ahead: at least one test for transitivity, "jump(a+b) = jump(a) followed by jump(b)"

<!-- grade_id: b_4_test, points: 1, grade: true, solution: true -->
```python
def test_transitive(a, b):
    """
    at least test for transitivity, jump(a+b) = jump(a) followed by jump(b)
    """
    # YOUR CODE HERE
    raise NotImplementedError()
    return  final_number == final_number2
    
assert test_transitive(10, 20)
assert test_transitive(999, 1234567)
```

B.5. getrandbits: at least one test to check exactly k bits are given

<!-- grade_id: b_5_test, points: 1, grade: true, solution: true -->
```python
def test_getrandombits(k):
    """
    Test exactly k bits are given
    """
    # YOUR CODE HERE
    raise NotImplementedError()
    
    
test_getrandombits(1)
test_getrandombits(32)
```

B.6.1. Kolmogorov-Smirnov statistic for random()

<!-- grade_id: b_6_1_test, points: 1, grade: true, solution: true -->
```python
import numpy as np
from statsmodels.distributions.empirical_distribution import ECDF
import matplotlib.pyplot as plt
%matplotlib inline

def test_uniform_KS():
    """
    Kolmogorov-Smirnov statistic for random()
    make sure the KS test statistics is small enough
    """
    N = 10000
    # YOUR CODE HERE
    raise NotImplementedError()
    
    # Kolmogorov-Smirnov test statistics, 95% confidence interval
    print(D_KS)
    assert D_KS < 1.36/np.sqrt(N)
    
    
test_uniform_KS()
```

B.6.2. binomial test for the single-bit frequency for getrandbits().

<!-- grade_id: b_6_2_test, points: 1, grade: true, solution: true -->
```python
import scipy.stats

def test_uniform_binomial():
    """
    binomial test for the single-bit frequency for getrandbits().
    make sure the p-value is large enough. 
    """
    N = 100000
    k = 10
    # YOUR CODE HERE
    raise NotImplementedError()
    assert p_value > 0.8
    
    
test_uniform_binomial()
```

B.7. tests that your class inherited the expected methods from the parent class

<!-- grade_id: b_7_test, points: 1, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

---
## Part c (2 points)
Use the Jupyter “magic” %timeit command to compare the amount of time it takes to generate $10^5$ PRNs using your SHA-256 method and using the default Python PRNG. Use the first cell to time your method, the second cell to time Python method, the third cell to discuss possible reasons for the difference. 

<!-- grade_id: c_1, points: 0.5, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: c_2, points: 0.5, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: c_3, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE